Human receptive endometrial organoid for deciphering the implantation window

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public Review):

Q1: First of all, the term organoid must be discarded. The authors just seed the endometrial cell mixture which assembles and aggregates into a 3D structure which is then immediately used for analysis. Organoids grow from tissue stem cells and must be passage-able (see their own description in lines 69-71). So, the term organoid must be removed everywhere, to not confuse the organoid field. It is not shown that the whole 3D assembly is passageable, which would be very surprising given the fact that immune and stromal cells do not grow in Matrigel because of the unfavorable growing conditions (which are targeted to epithelial cell growth).

We appreciate for your highlighting concerns regarding our organoid construction.

(1) The organoids in our system were originated from tissue stem cells.

We induced adult stem cells derived from endometrial tissue to construct organoids in vitro by various small molecules (such as Noggin, EGF, FGF2, WNT-3A and R-Spondin1), which involves a complex self-assembly process rather than a mere cellular assembly. Initially, there are single cells and small cell clusters in the system two days after the planting. On the fourth day, the glandular epithelial cells gradually assembled to glands, while the stromal cells spontaneously organized themselves around the glands. On the eleventh day, the endometrial glands enlarged, epithelial cells organized in a paving stone arrangement, and stromal cells established an extensive network. (Author response image1) (Figure 1C)

(2) The organoids we constructed are passage-able.

Most organoids were used for experiments up to the fifth generation, while some are extended to the 10th generation and cryopreserved. (Response Figure 1B, C)

(3) Immune and stromal cells are present in our system from the primary to the fourth generation. In our study, immune and stromal cells were identified not only from scRNA-seq data (third generation of organoids) (Figure 2A), but also from the morphology using 3D transparent staining and light sheet microscopy imaging (third generation of organoids), with Vimentin marking stromal cells, CD45 designating immune cells, and FOXA2 identifying glands. Further, flow cytometric analysis was applied to verify immune cells within the organoids (third generation of organoids). (Response Figure 1D, E, F)

Moreover, Immune cells and stromal cells can grow in Matrigel, which was also found in the study of organoid pioneer Hans Clevers (Hans Clevers et al., Nature Reviews Immunology 2019).

(A) The growth condition of endometrial cells was observed from day2 to day11 after plating under an inverted microscope.

Scale bar = 200 μm. (B) The endometrial organoids of different passages were observed from P1 to P5. Scale bar = 200 μm. (C) Stromal cells formed an extensive network (down). The arrowhead indicates dendritic stromal cells. Scale bar = 100 μm (left), Scale bar = 50 μm (right). (D) Exhibition of stromal cells marked by vimentin. Nuclei were counterstained with DAPI. The arrow indicates stromal cells. Scale bar = 40 μm (up), Scale bar = 30 μm (down). (E)Exhibition of immune cells marked by CD45 and endometrial gland marked by FOXA2. Nuclei were counterstained with DAPI. The arrow indicates immune cells. Scale bar = 50 μm. (F) Flow cytometric analysis of T cells and macrophages in the endometrial organoid. Gating strategy used for determining white blood cells (CD45+ cells), T cells (CD45+CD3+ cells) and macrophages (CD45+CD68+CD11b+ cells).

Q2: Second, the study remains fully descriptive, bombing the reader with a mass of bioinformatic analyses without clear descriptions and take-home messages. The paper is very dense, meaning readers may give up. Moreover, functional validation, except for morphological and immunostaining analyses (which are posed as "functional" but actually are only again expression) is missing, such as in vivo functionality (after transplantation e.g.) and embryo interaction. Importantly, the 3D structure misses the right architecture with a lining luminal epithelium which is present in the receptive endometrium in vivo and needed as the first contact site with the embryo. So, in contrast to what the authors claim, this is not the best model to study embryo interaction, or the closest model to the in vivo state (line 318, line 326).

Thank you.

(1) We have made the following improvements. Firstly, we have conducted additional experiments to validate the bioinformatics analysis. Secondly, the structure of the manuscript has been refined to ensure logical coherence and clear transitions between paragraphs. Thirdly, important findings have been emphasized to ensure readers’ comprehension and inspiration. Furthermore, the manuscript was revised by both domestic and international experts to enhance the readability and clarity.

(2) For the functional validation, in vivo transfer could not be carried out so far due to ethical limitation. But human embryos are able to develop and grow more efficiently in combining with the receptive endometrial organoids we generated (unpublished data).

(3) As you suggested, we replaced the “closest” with “closer”. It is undeniable that the model cannot completely simulate the in vivo implantation process that the luminal epithelium of the endometrium contacts the embryo first.

Q3: Third, receptive endometrial organoids (assembloids; Rawlings et al., eLife 2021) and receptive organoid-derived "open-faced endometrial layer" (Kagawa et al., Nature 2022) have already been described, which is in contrast to what the authors claim in several places that "they are the first" (e.g. lines 87-88, 316-319, etc). These studies used real organoids to achieve their model (and even showed embryo interaction), while in the present study, different cell types are just seeded and assembled. Hence, logically, immune cells are present which are never found in real organoid models. The only original aspect in the present study is the use of hormones to enhance the WOI phenotype. However, crucial information on this original aspect is missing such as concentration of the hormones, refreshment schedule, all 3 hormones added together or separately, and all 3 required?

Thank you for pointing out these researches referring to endometrial organoids.

(1) While we didn’t explicitly state "the first", we should be careful to use the expressions similar to "the first". It has been changed to a gentle and modest expression, as follows “we are far from understanding how embryo implantation occurs during the WOI due to ethical limitations and fewer in vitro receptive endometrial model” and “which confirms that they are closer to the in vivo state”.

(2) The definition of organoids and the existence of immune cells have been detailed addressed in the first question.

(3) In terms of hormone scheme, hormone concentrations have been detailed in Table S2 of Supplementary. Estrogen was supplemented to the basal medium for the initial two days, after which a combination treatment of MPA, cAMP, PRL, hPL, and HCG was administered for the subsequent six days. The medium was refreshed every two days.

All three hormones were deemed necessary, which was validated by multiple group comparisons. Only the organoids treated with all six hormones together exhibited an endometrial receptivityrelated gene expression profile. (Author response image 2).

Heatmap showing receptivity related gene expression profile of organoids in each hormone regimen.

Q4: Moreover, it is not a "robust" model at all as the authors claim, given the variability of the initial cell mixture (varying from patient to patient). Actually, the reproducibility is not shown. The proportions of the different cell types seeded in the Matrigel droplet will be different with every endometrial biopsy. It would be much better to recombine epithelial (passageable) organoids with stromal and immune cells in a quantified, standardized manner to establish a "robust" model.

Thanks for your suggestion.

Firstly, the constructed endometrial organoids generally consist of epithelial, stromal, and immune cells. However, it is undeniable that the cell proportions may vary slightly among different patients. Secondly, the term "robust" is intended to convey strong support for embryo development, which will be supported by our next study (unpublished data). Therefore, robust is replaced here as alternative. Thirdly, as for "reproducibility", the hormone-treated organoids from different women exhibited similarity to the in vivo receptive endometrium through multi-omics analysis, ERT, and various other experiments.

Reviewer #2 (Public Review):

Q1: With endometrial receptivity analysis, they suggest a successful formation of the implantation window in vitro, but this result is difficult to interpret.

Thanks for your question.

We understand that the most effective way to demonstrate endometrial receptivity is embryo implantation, which was conducted simultaneously and will be presented in our next study. In this study, we validated the receptivity based on the current researches.

(1) At the single-cell transcriptome level, the cellular composition and function of the receptive endometrial organoids were similar to those of the in vivo implantation window (Stephen R. Quake et al, 2020).

(2) At the whole organoids level, the receptive endometrial organoids exhibited the similar characteristics in transcriptome and proteome to the in vivo mid-secretory endometrium (Andres Salumets 2017, Qi Yu 2018, Triin Laisk 2018, Edson Guimarães Lo Turco 2018, Xiaoyan Chen 2020, Francisco Domínguez 2020, DavidW. Greening 2021, Norihiro Sugino 2023). The receptive endometrial organoids were also validated by endometrial receptivity test (ERT), which utilized high-throughput sequencing and machine learning to assess endometrial receptivity (Yanping Li et al., 2021).

(3) At the microstructural level under electron microscope, the receptive endometrial organoids exhibited characteristics of the implantation window, such as pinopodes, glycogen particles, microvilli, and cilia.

Overall, the receptive organoids we constructed closely resemble the in vivo implantation window at the single-cell, organoids, and microstructural levels based on existing researches.

Q2: Analyzing transcriptome and proteome information of WOI organoids, authors demonstrate a strong response to estrogen and progesterone, but some comparisons are made with CTRL and SEC, and others only with CTRL, which limits the power of some results. In the same way, some genes related to Cilia and pinopodes appear dominant in WOI organoids, but the comparison by electron microscopy is made only against CTRL organoids.

In subsequent analysis, WOI organoids showed a marked differentiation from proliferative to secretory epithelium, and from proliferative epithelium to EMT-derived stromal cells than SEC organoids. These statements are based on their upregulation of monocarboxylic acid and lipid metabolism, their enhanced peptide metabolism and mitochondrial energy metabolism, or their pseudotime trajectories. However, other analyses (such as the accumulation of secretory epithelium or decreased proliferative epithelium, the increased ciliated epithelium after hormonal treatment, or the presence of EMT-derived stromal cells) show only small differences between SEC and WOI organoids.

Thank you for raising these important questions.

(1) At the organoid level, the differences in transcriptome and proteome between SEC and WOI organoids are not significant. This is understandable because WOI organoids are further induced towards the implantation window based on the secretory phase (i.e. SEC organoids), and both are similar at the overall organoid level.

(2) At the single-cell level, the accumulation of secretory epithelium, decreased proliferative epithelium, increased ciliated epithelium post hormonal treatment, or the presence of EMTderived stromal cells are the fundamental features of the secretory endometrium. Therefore, these features are present in both WOI and SEC organoids. However, the most notable differences lie in the more comprehensive differentiation and varied cellular functions exhibited by WOI organoids compared to SEC organoids.

(3) Regarding electron microscopy, we have now quantitatively compared the presence of various characteristic structures such as microvilli, cilia, pinopodes and glycogen in the CTRL, SEC and WOI groups. It has been observed that WOI organoids possess longer microvilli and increased cilia, glycogen, and pinopodes compared to SEC organoids (Fig2H).

Reviewer #1 (Recommendations For The Authors):

Q1: Several of the key methods are performed by companies, hence not in detail described and therefore not verifiable which is essential for reviewers and readers.

We are grateful for the suggestion. Specific methods have now been incorporated into the "Supporting Information" section. (Line91~102, Line 107~123, Line 132~139)

Q2 - Line 49: It is not shown in the present study whether the WOI organoids are a 'robust' platform.

- Line 76: There is a study (Dolat L., Valdivia RH., Journal of Cell Science, 2021) that developed a co-culture with endometrial organoids and immune cells (neutrophils) which should be mentioned.:

We have reweighed the word and now replace 'robust' with 'alternative' (Line 54). We have considered the reviewer's suggestion and added this citation (Line 82-83) about the cocultivation of immune cells with endothelial organoids, which was not previously cited mainly because the research model was mouse.

Q3: Figure 1: Endometrial organoids possess endometrial morphology and function. - The authors should further explain their decision to add PRL, hCG, and hPL to the organoid culture. Why these particular compounds? What is their specific role during the WOI?

In terms of hormone scheme, estrogen and progesterone promote the transition of endometrial organoids into the secretory phase, and on this basis, pregnancy hormones can further promote their differentiation. PRL promotes immune regulation and angiogenesis during implantation, HCG improves endometrial thickness and receptivity, and HPL promotes the development and function of endometrial glands. Our constructed WOI organoid is in a state conducive to embryo implantation. We aim to develop an in vitro model for embryo implantation study. The detailed explanation of this aspect was initially provided in the Discussion section (Lines 298–313). To enhance the clarity for reviewers and readers regarding the selection of the hormonal regimen, we have now articulated it in the Results section (Lines 124–130).

When selecting hormone formulations, multiple group comparisons were made. It was found that the number, area, and average intensity of organoids in these groups were similar over time. But the WOI organoids showed endometrial receptivity related gene expression profile, which highly expressed genes positively correlated with endometrial receptivity, and lowly expressed genes negatively correlated with receptivity, compared to the other hormone formulations (added to Figure S1E, S1F). Hormone dosage was primarily based on peri-pregnant maternal body or localized endometrium levels (Margherita Y. Turco et al., Nature Cell Biology 2017).

- Line 108: "the endometrial cells" instead of "endometrial organoid"? Because the authors also refer to the stromal cells.

You should be referring to this sentence “The endometrial organoid, consisting of vesicle-like glands, fibrous stromal cells, and other surrounding cells, developed into a 3D structure with the support of Matrigel”. Organoid, a self-assembled 3D structure, consists of multiple cells and closely resembles in vivo tissue or organ. It offers high expansibility, phenotypic, and functional properties. Here, we aim to delineate the endometrial organoid, comprising epithelial cells, stromal cells, and other cellular components that assemble to form intricate 3D structures. Hence, the term "endometrial organoid" is more appropriate.

- Line 110: "the endometrial glands", do the authors mean the endometrial organoids? The authors also mention they enlarge, which must be quantified.

You should be referring to this sentence “As the organoids grew and differentiated, the endometrial glands enlarged, epithelial cells adopted a paving stone arrangement, and stromal cells formed an extensive network”. Here, we mean the “endometrial glands” grow progressively in the organoids. We agree with your suggestion to quantify the change of organoids’ area over time, and found that they increased progressively in all three groups (shown as follows) (Fig.S1E) (Line130-131)

The dynamic changes of the area of organoids over time in the CTRL, SEC and WOI organoids.

- Line 112: E-cadherin is a general epithelial marker, not a glandular marker.

We agree with your suggestion and now change to ‘The epithelium marker E-cadherin’ (Line110).

- Line 116: Which group was used for KI67 and CC3 staining?

The CTRL organoids were used for Ki67 and CC3 staining. We have modified this expression in the Figure 1E Legend.

- Line 123: Organoid size (diameter or area) needs to be quantified to claim that WOI organoids grow slower than SEC/CTRL organoids. The same goes for Ki67+ cells for proliferation. In the legend of Fig 1B, the authors in contrast state that the organoids show a similar growth pattern.

We are extremely grateful to you for pointing out this problem. We quantitatively analyzed the size of organoids in the three groups. The area was found to be increasing over time, with the three groups growing the most vigorously in the CTRL group, followed by the SEC group and the WOI group, but the differences were not statistically significant. Relevant results have been added to Figure S1E (Line130-131). There were no significant differences in Ki67 expression of these organoids. Therefore, the three groups of organoids showed a similar growth pattern. We decided to delete the statement “Following hormonal stimulation, WOI organoids exhibited slower growth than SEC and CTRL organoids, while CTRL organoids maintained robust proliferative activity (Fig. 1B)”.

The dynamic changes of the area of organoids over time in the CTRL, SEC and WOI organoids.

- Line 126: Fourteen days of organoid treatment is a very long time. Growing organoids may already be dying which should be checked by CC3 staining to prove that organoids are still fully viable.

Endometrial organoids are vigorous in proliferation and have a long survival period due to the presence of adult stem cells. To address your queries effectively, we conducted CC3 staining on the organoids treated for 14 days, revealing negligible expression levels (shown as below).

Figure note: The Ki67 and CC3 immunostaining on the organoids after 14-day hormone treatment.

- Line 128: Changes in hormone receptors should be supported by RT-qPCR data to be more convincing

We agree with your suggestion. Here we supplemented the RT-PCR results of hormone receptors as follows (Figure S1D) (Line119-121). PAEP and PGR are associated with progesterone, and OLFM4 and EGR1 are associated with estrogen.

- 1A: Are authors able to see and characterize decidualized stromal cells as indicated in the illustration?

Upon the reviewer's inquiry, we carefully observed the morphology of stromal cells in hormone-treated organoids. Regrettably, the morphology of decidualized stromal cells was not ascertainable through light microscopy in our endometrial organoids.

- 1C: Which treatment condition are the organoids in these images?

This figure showed the bright-field morphology of the CTRL organoids, which is now noted in the Figure 1C legend.

- 1D: PAS staining should be quantified to support the claims.

We agree with your suggestion. The quantitative comparison of PAS staining was conducted in these three groups of organoids (Figure S1G) (Line142-143)

- 1D: Where are the stromal cells in the model? There should be vimentin-positive cells outside of the glands.

The figure 1D illustrates the outcomes of section staining, which owned limitation to displaying stromal cells around the gland. Considering the 3D structure of organoids, we conducted organoid clearing and staining, and observed stromal cells (marked by Vimentin) under light sheet microscope (shown as below). The stromal cells were also presented using this method in the original Figure 2B.

Exhibition of stromal cell marked by vimentin of CTRL organoid through whole-mount clearing, immunostaining and light sheet microscopy imaging.

Nuclei were counterstained with DAPI. The arrowhead indicates stromal cells. Scale bar = 70 μm.

Figure 2: Developing receptive endometrial organoids in vitro mimicking the implantation window endometrium.

- Line 142: CD44 is not an exclusive marker for immune cells. It has been shown to be expressed in glandular secretory epithelial cells (Fonseca et al., 2023). The authors also mention that CD44 is expressed in stromal cells (line 265). Staining for CD45 (or another immune-specific marker) is needed to demonstrate the presence of immune cells.

We appreciated your suggestions. We demonstrated the distribution of immune cells in organoids using the organoid clearing technique in combination with light-sheet microscopy imaging, using CD45 as a marker (Figure 2C).

- Line 144: What are the proportions of the immune cells? What is the variation between patient samples?

We assessed the proportion of immune cells with the help of flow cytometry and analyzed the proportion of Macrophages and T cells in organoids derived from 8 patients. The proportion of WBC in organoids was about 3%~4% (Figure 2D), among which macrophages were less than 1% and T cells less than 2% (Figure S2E). There existed a very few patients with large heterogeneity, and the proportion of immune cells in most patients was

relatively stable.

- Line 161: What is the endometrial receptivity test (ERT)? Not explained at all.

Endometrial Receptivity Test (ERT) is a kind of gene analysis-based method for detecting endometrial receptivity, which combines high-throughput sequencing and machine learning to analyze the expression of endometrial receptivity-related genes, allowing for a relatively accurate assessment of endometrial receptivity. It is currently used in clinical practice to determine endometrial receptivity and guide personalized embryo transfer (Yanping Li et al., J Transl Med 2021). (line179-183)

- 2A: The authors' dataset is compared to a published dataset. How were they combined? Were they merged, mapped on each other, or integrated? Were all cells employed from the published dataset or specific cell types? Much detail to evaluate the analysis is missing.

We are very grateful for your comments.

(1) The four raw datasets (CTRL, SEC and WOI organoids, and mid-secretory endometrium) underwent batch correction and integration using Harmony. Subsequently, the integrated dataset underwent dimensionality reduction via PCA. The soft k-means clustering algorithm was employed to address batch effects and clustering, utilizing a clustering parameter resolution of 0.5. Finally, the clustering results were visualized using tSNE based on the cell subpopulation classification. (“Methods” Line164-175)

(2) The Figure 2A displayed comparison of glandular and luminal epithelium, secretory epithelium, LGR5 epithelium, EMT-derived stromal cells, ciliated epithelium, and glandular secretory epithelium (shown as Figure S2C~S2D) (Line150-154)

- 2E: Please add the cell type names above the heatmaps to improve readability.

Thanks to your suggestion, we have added the cell type names above the heatmaps.

- 2G: The difference between the left and right graphs is not clear from the figure itself. Improve by adding a title and more explanation.

Thanks for your careful review. We have added the title to the left and right graphs.

Supplementary Figure 3 is referenced with Figure 2. Supplementary Figure 2 is referenced with Figure 3. The order needs to be changed.

Thanks for your careful review. We have changed the order.

- S3B: Typical markers for annotation of the different cell clusters are not included and therefore it is not convincing enough that annotations are correct. E.g. Epithelial markers (EPCAM, CDH1), Stromal cells (VIM, PDGFRA), SOX9+LGR5+ cells (SOX9, LGR5). How were the EMT-derived stromal cells designated? It is not clear from the data whether they are in fact EMT-derived or whether they show epithelial markers as well (stated in line 246).

We deeply appreciate your suggestion. We provided more details to describe the cell clustering as the following. Single-cell transcriptomics analysis referred to CellMarker, PanglaoDB, Human Cell Atlas, Human Cell Landscape, and scRNASeqDB, and previous endometrium related studies. (W. Wang et al., Nat Med 2020, P. D. Harriet C. Fitzgerald et al., PNAS 2019, K. M. Thomas, M Rawlings et al., eLife 2021, L. Garcia-Alonso et al., Nat Genet 2021)

(1) SOX9+LGR5+ cells: SOX9 and LGR5 are both proliferative markers. SOX9 is expressed in all clusters dispersedly. LGR5 is mainly expressed in two clusters, one of which is stem derived epithelium, and the other cluster expresses LGR5 in a scattered manner. Refer to the markers of SOX9+LGR5+ cells, SOX9+LGR5- cells, and SOX9+ proliferative cells in 2021 Nature Genetics (L. Garcia-Alonso et al., Nat Genet 2021), the cells in this cluster expressed high levels of NUAK2, CNKSR3, FOS and LIF, which was consistent with the expression profiles of SOX9+LGR5+ cells and SOX9+ proliferative cells. However, considering that the number of cells expressing LGR5 was relatively small, this cluster of cells was renamed SOX9+ proliferative epithelium.

Figure 3: Receptive endometrial organoids recapitulate WOI-associated biological characteristics. - Line 173-174: The WOI organoids should be compared in detail to the SEC organoids in addition to the CTRL organoids, to show that this WOI model and new hormonal treatment is providing better results compared to the SEC organoids and the results obtained in previous studies.

Thanks for your suggestion. At the organoid level, the differences in transcriptome and proteome between SEC and WOI organoids are not significant. This is understandable because WOI organoids are further induced towards the implantation window based on the secretory phase (i.e. SEC organoids), which prompted us to continue exploring at the single-cell level.

- Line 190: Quantification of pinopodes is required to claim that they are more densely arranged in WOI organoids.

- Line 190-191: Again, is there a difference in pinopode presence between the WOI and SEC organoids to show that the WOI organoids are really distinct and a better model?

We agree with the reviewer’s suggestion and quantified the pinopodes. The CTRL, SEC and WOI organoids were found to have increasing numbers of pinopodes, with WOI organoid owning the most abundant pinopodes under electron microscope. (Figure 2H) (Line184-186)

- Line 194: Also here, quantification of the glycogen particles is missing.

We agree with your suggestion. We have quantified the area of glycogen particles under electron microscope in the CTRL, SEC and WOI organoids. It was found that WOI organoid had the most glycogen particles. (Figure 2H) (Line184-186)

- 3C: There is no difference between SEC and WOI organoids condition for OLFM4 and PRA/B. What is the purpose then of adding extra hormones if no difference is present?

The figure 3C indicated that there was no significant difference in OLFM4 and PRA/B level (reflecting estrogen and progesterone responsiveness) in SEC and WOI organoids at the organoids level. It is understandable because WOI organoids are induced further into the implantation window on the basis of the secretory phase (i.e., SEC organoids), and both are similar at the overall level of organoids. Based on this, we further explored the differences between WOI organoids and SEC organoids at the single-cell level.

- 3G: A higher magnification is necessary to evaluate cilia staining. From these images, it seems like CTRL organoids also express acetyl-a-tubulin.

Thanks for your suggestion. The figure has been enlarged and shown as below. The acetyl-a-tubulin of WOI organoids is different from that of CTRL organoids in morphology and expression level. The glands of WOI organoids have small green tips (expressing acetyl-α-tubulin) convex toward the lumen. WOI organoids expressed higher level of acetyl-α-tubulin than CTRL organoids. (Now replaced with Figure 3G in the revised draft).

Figure 4: Structural cells construct WOI with functionally dynamic changes

- Line 211: To which figure are these claims referring to?

You should be referring to this sentence “In terms of energy metabolism, the WOI organoids exhibited upregulation of monocarboxylic acid and lipid metabolism, and hypoxia response”. Up-regulation of monocarboxylic acid and lipid metabolism in WOI organoids is reflected in Figure 3B, and up-regulation of hypoxia responses is reflected in Figure S3F.

- In general, it should be stated in the text that CellPhoneDB is a useful tool to investigate ligandreceptor interactions, however, it only proposes potential interactions. To validate such interactions, stainings and functional assays are required.

Thanks for your suggestion. The CellphoneDB was briefly introduced in the "Methods" section of "Supporting information" originally. Now it has been explained in the line 256-257 of main text.

We agree that staining and functional assays are required to validate the ligand-receptor interactions. Therefore, we used the proximity ligation assay (PLA) to verify the trend of interaction. (Figure S2J, Line259-261, Line 277-279, Line 285-288）

- Line 243: Please describe the process of EMT in the endometrium more specifically.

EMT is a common and crucial biological event in the endometrium during the implantation window. During the EMT process, epithelial cells lose their epithelial characteristics while gaining migratory and invasive properties of fibroblasts.

During the attachment and adhesion phases of embryo implantation, interaction mediated by trophoblastic factors (e.g. integrins) and maternal ECM factors (e.g. fibronectin) induce the eventual EMT in the trophectoderm. During the peri-implantation period, microRNAs, (e.g. miR429 and miR-126a-3p) which regulate EMT, are expressed in the maternal luminal epithelium to different degrees, mediating its transformation process as the blastocyst invades the maternal decidua. The epithelium of endometrium transforms to epithelioid stromal cells with increased migratory and invasive capacities through the EMT process. The decidual stromal cells migrate away from the implantation site, having acquired increased motility. (Line 265-267)

- Lines 247-251 and 313-316: the claim that proliferative epithelium transforms into EMT derived stromal cells by pseudotime trajectory is too bold and must be underpinned by other means. Pseudotime analysis only suggests and is by definition biased since the first/originating population must be defined by the operator.

In addition to pseudotime analysis based on monocle, RNA rate analysis based on scVelo is also used for cell evolution analysis. They can prove each other if both analyses indicate the transformation from proliferative epithelium to EMT-derived stromal cell. RNA rate analysis automatically determines the direction of differentiation, which can be used as evidence to determine the starting point of pseudotime analysis.

RNA rate analysis showed that the EMT derived stromal cell was most closely connected to the proliferative epithelium. Besides, the pseudotime point plot inferred that the proliferative epithelium was the root cell. It can be mutually proved with pseudotime analysis that the transformation from proliferative epithelium to EMT-derived stromal cell.

RNA rate junction diagram (To infer intercellular connectivity)

Time differentiation of cells

Discussion

- Line 300-302: It would be interesting to investigate ATP production and IL8 release in the WOI organoids to validate with findings from in vivo.

To answer this point of your interest, we purposely examined ATP production and IL8 release. It was found that WOI organoids indeed produced much more ATP and IL8 than CTRL and

SEC organoids (Figure S3L) (Line323-324)

- Line 313-316: Do the WOI organoids lose polarity and cell-to-cell junctions?

Transcriptome sequencing revealed downregulation of cell adhesion and RHO GTPase signaling in WOI organoids (Figure 3B). Electron microscopy revealed that the cellular arrangement of WOI organoids was slightly looser than that of CTRL organoids, but the microvilli were still oriented toward the medial side of the glands and did not undergo polarity reversal (shown as below).

Electron micrograph of the CTRL (left), and WOI (right) endometrial organoid.

Scale bar = 5 μm.

- Line 322: Where is the data that shows that 'a decreased abundance of immune cells', is observed?

A decreased abundance of immune cells was observed through single-cell transcriptome sequencing and flow cytometry. The number of immune cells was reduced in WOI organoids compared to CTRL organoids in single-cell sequencing results (Figure 4A). Besides, flow cytometry also showed that the percentage of WBCs in WOI organoids was lower than that in CTRL organoids (Figure S2F).

- Line 324: Elaborate more on how the immune cell composition differs from the endometrium.

The differences of immune cell composition between organoids and endometrium were mainly reflected in the proportion of WBC, the proportion of immune cell subtypes and the changes of T cells after entering the implantation window.

Firstly, the proportion of WBCs in organoids was lower than that in endometrium. Flow cytometry showed that the proportion of WBC in organoids was about 3%~4% (Figure 2D), but the proportion of WBCs in endometrium was about 8% (W. Wang et al., Nat Med 2020). Secondly, the proportions of T cells and macrophages in organoids were about 2%~3% and 1% (Figure 2D), respectively, but the proportions of lymphocytes and macrophages in endometrium were 7%~8% and 0.6%~0.7% (W. Wang et al., Nat Med 2020). Besides, after entering the implantation window, T cells in WOI organoids decreased (Figure S2F), while T cells in endometrium increased (W. Wang et al., Nat Med 2020). These three aspects have differences in vivo and in vitro. (Line347353)

Material and Methods

- What are the concentrations of all medium components?

Thanks to your suggestions. The concentrations of all medium components have now been refined in Table S1.

- Authors mention 10x while Smartseq2 is mentioned in Dataset S7?

Thanks for your careful review. Single cell transcriptome sequencing in this study was done using 10X Genomics. Smartseq2 was used to sequence the transcriptome of a gland and its surrounding cells, which can be regarded as small bulk RNA sequencing. A small number of cells are utilized in Smartseq2 to construct a full-length mRNA library with enhanced transcript sequencing coverage, making it particularly well-suited for small-scale samples such as organoids.

The data in Dataset S7 are acquired from small bulk RNA-seq with Smartseq2.

Reviewer #2 (Recommendations For The Authors):

Q1: The theoretical choice of extra reagents added to the WOI organoids culture (PRL, hCG, and hPL) is theoretically justified, but not experimentally. On what previous studies, or performed experiments, are the choice of conditions used based?

When selecting hormone formulations, multiple group comparisons were made. It was found that the number, area, and average intensity of organoids in these groups were similar over time. But the WOI organoids showed endometrial receptivity related gene expression profile, which highly expressed genes positively correlated with endometrial receptivity, and lowly expressed genes negatively correlated with receptivity, compared to the other hormone formulations (added to Figure S1E, S1F). Hormone dosage was primarily based on peri-pregnant maternal body or localized endometrium levels (Margherita Y. Turco et al., Nature Cell Biology 2017).

Q2: Text in line 111 indicates that "stromal cells formed an extensive network", but vimentin fluorescence is not present on any image surrounding organoids in that figure. This assertion could only be supported by the subsequent results in Figure 2B. In addition, it is not indicated what kind of organoids have been used for these experiments

The stromal cells arranged around the glands in the 3D structure (as shown in Figure 1C and Figure 2B), where bright-field high magnification photography, clearing staining of the organoids, and light microscopy imaging were used, respectively. However, there are many steps of fixation, embedding, staining and elution during the immunostaining of sections. It is difficult to preserve the arrangement and morphology of the stromal cells in the slice, so the stromal cells were not intentionally captured in the other images.

Figure 1C and Figure 2B are both CTRL organoids, which are now noted in the corresponding figure legend section.

Q3: It is not clear how glycogen secretion into the lumen is assessed in Figure 1D.

Glycogen from the subnuclear region of the glandular cells gradually reaches the top of the cells, i.e., the supranuclear region, and is discharged into the glandular lumen as parietal plasma secretion. Glycogen-containing eosinophilic secretion can be seen in the glandular lumen in Figure1D.

Q4: Assertions about differences in proliferation between groups are purely subjective; some kind of measurement and analysis would be necessary to be sure that there is differential proliferation based on Figure 1B.

We are extremely grateful to you for pointing out this problem. We quantitatively analyzed the size of organoids in the three groups. The area was found to be increasing over time, with the three groups growing the most vigorously in the CTRL group, followed by the SEC group and the WOI group, but the differences were not statistically significant. Relevant results have been added to Figure S1E (Line130-131).

Q5: For progesterone receptor expression analysis organoids are cultured for fourteen days. What is the basis for this change in culture time?

The choice of time point here is based on the secretary period of 14 days in the female menstrual cycle, when the endometrium is stimulated by estrogen and progesterone to maximized

level.

Q6: "n" number of individuals analysed through single-cell transcriptomics is not indicated.

One patient's endometrium was simultaneously constructed into CTRL, SEC and WOI organoids, which were then subjected to single-cell transcriptome sequencing. This is described in the Supporting Information (Line 141-142).

Q7: Where does the classification of EMT-derived stromal cells come from?

EMT is a common and crucial biological event in the endometrium during the implantation window. During the EMT process, epithelial cells lose their epithelial characteristics while gaining migratory and invasive properties of fibroblasts.

This cluster of cells expresses both epithelium markers CDH1 and EPCAM, and specifically expresses high levels of the EMT-related stromal cell markers AURKB, HJURP and UBE2C. During endometrial EMT, AURKB upregulates MMP2, VEGFA/Akt/mTOR and Wnt/β-catenin/Myc pathways to induce EMT (Zhen Wang et al., Cancer Manag Res 2020). HJURP also activates Wnt/β-catenin signaling to promote EMT (Y Wei et al., Eur Rev Med Pharmacol Sci 2019, Tianchi Chen et al., Int J Biol Sci 2019). UBE2C is upregulated by estrogen to promote EMT (Yan Liu et al., Mol Cancer Res 2020). Therefore, this cluster was defined as "EMT-derived stromal cells”.

Q8: In the endometrial receptivity test (ERT), endometrium sample data matches with prereceptive endometrium and WOI organoids data matches with a receptive endometrium, but why there is no information about CTRL and SEC organoids?

We performed ERT on these samples at a time when our hospital has a cooperative project with Yikon Genomics (Jiangsu, China). However, only endometrium and WOI organoids were sent for testing due to the limited quotas. Considering the end of cooperation and batch effect, no more CTRL and SEC organoids were tested. Moreover, the current ERT is a machine learning model based on the sequencing data of endometrium samples. But there are still differences in cellular composition between endometrial organoids and endometrium. Thus, the results need to be interpreted in conjunction with other results.

Q9: When analysing the transcriptome and proteome, some comparisons are made between WOI vs CTRL and SEC, or just WOI vs CTRL. It would be interesting to have all the comparisons since the power of WOI organoids lies in their differences with SEC organoids.

Thanks for your suggestion. At the organoid level, the differences in transcriptome and proteome between SEC and WOI organoids are not significant. This is understandable because WOI organoids are further induced towards the implantation window based on the secretory phase (i.e. SEC organoids), which prompted us to continue exploring at the single-cell level.

Q10: Electron microscopy comparisons with respect to pinopods, cilia, and microvilli are only performed between WOI and CTRL. It would be interesting to check it with SEC.

We now quantitatively compared the presence of various characteristic structure like microvilli, cilia, pinopodes and glycogen in the CTRL, SEC and WOI organoids. It was found that WOI organoid had longer microvilli and increased cilia, glycogen, and pinopodes (Figure 2H).

Q11: Line 190 states that pinopods are arranged more densely in WOI organoids than in CTRL organoids. Seems to be a subjective observation. Is there an objective method to quantify this?

We agree with the reviewer’s suggestion and quantified the pinopodes. The CTRL, SEC and WOI organoids were found to have increasing numbers of pinopodes, with WOI organoid owning the most abundant pinopodes. (Figure 2H) (Line184-186)

Q12: Some characteristics are very similar between WOI and SEC organoids (such as the accumulation of secretory epithelium or decreased proliferative epithelium, the increased ciliated epithelium after hormonal treatment, or the presence of EMT-derived stromal cells). The authors should complement the discussion by objectively justifying the use of WOI versus SEC organoids. Would they be useful in more specific cases or at a general level when studying implementation?

Thanks for your comments. WOI organoids are differentiated from SEC organoids towards the implantation window. Therefore, WOI organoids are suitable for studying periimplantation physiological changes or exploring pathological mechanisms. SEC organoids can be used when studying only a range of pathological problems such as endometrial secretory phase changes or hormone reactivity. (Line 365-368)

Q13：ExM media is described in Table S1, but it does not include the concentration of the different reagents in the culture medium, which is the most interesting data about the ExM medium.

Thanks to your suggestions. The concentrations of all medium components have now been refined in Table S1.

Q14: It is not specified which organoid pass is used in each experiment. Is it always the same pass?

Our experiments were conducted using P1~P3 generation endometrial organoids, as specified in the “Supporting Information” Line 54~55.

Q15: As a protocol for freezing organoids is included in materials and methods, do the authors use freshly cultured organoids or do they cryopreserve them and thaw them for culturing?

Thanks for your question. We used freshly cultured organoids in the manuscript. We listed the freezing protocol to illustrate that the constructed organoids can be frozen and recovered for special experimental needs and the establishment of sample banks.

Q16: The most important point: Neither of the two studies that developed human endometrial organoids from tissue biopsies (Boretto et al. 2017 and Turco et al. 2017), observed stromal cell growth in culture. They disappeared between the first and second pass (as indicated by Turco et al. 2017). How do the authors justify the presence of stromal cells in their organoid culture if they rely on the protocols previously described by these research groups? If it is the case that they can only use the initial pass (freshly planted cells from endometrium), it does not make sense to include the freezing of the different passes in materials and methods, since the expansion capacity of the culture would be lost, which implies a major limitation of the model.

Thanks for your question.

(1) We did not completely follow the protocols of these research groups. To maximize the recovery of both epithelial and stromal cells, we optimized key steps such as tissue digestion and cell strainer filtration. We shortened the digestion time to 20 minutes to protect cells from the digestion solution and retain some cell aggregates, which are beneficial for maintaining cell stemness and preserving stromal and immune cells cluster. The 40 μm filter membrane was used to isolate the endometrial cells, which may acquire both epithelial, and stromal cells.

(2) Our experiments were conducted using P1~P3 generation of freshly constructed organoids. However, we also used recovered organoids when fresh endometrial samples were not available due to the COVID-19 epidemic. It was found that the organoids (e.g., P0~P5) still exhibited vigorous growth condition after recovery and could continue to be cultured by passaging (shown as below).

The recovered organoids can be used for special experiments and biobank establishment.

The endometrial organoids of different passages were observed before cryopreservation and after recovery.

Scale bar = 200 μm.

Q17: It is not clear which organoids include Figure S2F. Does it include the three types of organoids or just WOI organoids?

This circle diagram showed the functions of upregulated genes in the WOI group compared to CTRL group from combined transcriptome and proteome analysis, which has been labeled in the figure legend section.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Human endometrial receptivity is a critical determinant of pregnancy success; however, in vivo studies of its features and regulation are particularly challenging due to ethical restriction. Recently, the development of human endometrial organoids has provided a powerful model to examine this intricate biological process. In this study, we established a specialized human window-of-implantation (WOI) endometrial organoid system that mimics the in vivo receptive endometrium. The endometrial organoid recapitulates not only the structural attributes of glands and stroma, but also the in vivo endometrium’s cell types at the single-cell level. Furthermore, the WOI endometrial organoid exhibits the features of in vivo implantation endometrium such as hormone responsiveness, secretory functions, and distinctive implantation window hallmarks, including decidualization, extracellular matrix (ECM) remodeling, pinopode formation, cilia generation and metabolism. This in vitro model also effectively demonstrates the process of proliferation-secretory transformation, ciliary epithelium differentiation, and epithelial-mesenchymal transition (EMT). As such, WOI organoids hold great promise as a robust platform to unravel the intricate mechanisms governing endometrial receptivity regulation, maternal-fetal interactions, and associated diseases, ultimately driving impactful advancements in the field.

Revised version: Human endometrial receptivity is a critical determinant of pregnancy success; however, in vivo studies of its features and regulation are particularly challenging due to ethical restriction. Recently, the development of human endometrial organoids has provided a powerful model to examine this intricate biological process. In this study, we established a specialized human window-of-implantation (WOI) endometrial organoid system that mimics the in vivo receptive endometrium. The endometrial organoid recapitulates not only the structural attributes of glands and stroma, but also the in vivo endometrium’s cell types at the single-cell level. Furthermore, the WOI endometrial organoid exhibits the features of in vivo implantation endometrium such as hormone responsiveness, secretory functions, and distinctive implantation window hallmarks, including decidualization, extracellular matrix (ECM) remodeling, pinopode formation, cilia generation and metabolism. This in vitro model also effectively demonstrates the process of proliferation-secretory transformation, ciliary epithelium differentiation, and epithelial-mesenchymal transition (EMT). As such, WOI organoids hold great promise as an alternative platform to unravel the intricate mechanisms governing endometrial receptivity regulation, maternal-fetal interactions, and associated diseases, ultimately driving impactful advancements in the field.

Introduction

The human endometrium, comprising various cell types, experiences shedding (menstrual phase), regeneration (proliferative phase), and differentiation (secretory phase) under the regulation of estrogen and progesterone during the menstrual cycle (1). A brief window, called the “window-of-implantation (WOI)” or mid-secretory phase, allows for embryo implantation into the endometrium (1). The endometrium at the embryo implantation site contains diverse cells, including ciliated epithelial cells, decidualized stromal cells and immune cells. Cyclic endometrial changes akin to those in humans are exclusive to apes, Old World monkeys, molossid bats, and spiny mice, but not mice (1), rendering typical mouse models insufficient for accurately simulating human endometrium. Existing endometrial cell lines, such as Ishikawa and HEEC cells, consist solely of a single cell type, thus failing to reproduce the intricate physiological structure and function of the endometrium.

Earlier version: Replicating and reconstructing human organs has become essential for exploring tissue physiology and function. Organoid, a self-assembled 3D structure, closely resemble in vivo tissue or organ (2). They offer high expansibility, phenotypic, and functional properties, emerging as powerful tools for investigating tissue physiology and disease (3). In 2017, the first long-term and hormone-responsive human endometrial organoid was established using adult stem cells from endometrial biopsy(2). Based on this, Margherita Y. Turco, et al developed noninvasive methods to construct endometrial organoids from menstrual flow(4). Apart from adult stem cells, pluripotent stem cells were also induced to endometrial stromal fibroblasts and epithelium, and then cocultured to form organoids(5, 6), which offered vigorous proliferative capacity but lacked immune cells and other components of the microenvironment. In efforts to mimic human endometrial cell types in vitro, the endometrial organoids progressed from epithelial organoids(7), to assemblies of epithelial and stromal cells(8, 9) and then to stem cell-laden 3D artificial endometrium(10, 11), which were solely closer but not completely identical to the endometrium. Additionally, pathological endometrial organoid models have been established for conditions like endometriosis, endometrial hyperplasia, Lynch syndrome, and endometrial cancer (12). These organoids simulated endometrial morphology, hormone responsiveness, and physiological and pathological processes in vitro (3, 7, 13–15), facilitating the study of physiological phenomena (13, 16, 17), pathogenic mechanisms(14) and drug screening(12).

Revised version: Replicating and reconstructing human organs has become essential for exploring tissue physiology and function. Organoid, a self-assembled 3D structure, closely resemble in vivo tissue or organ (2). They offer high expansibility, phenotypic, and functional properties, emerging as powerful tools for investigating tissue physiology and disease (3). In 2017, the first long-term and hormone-responsive human endometrial organoid was established using adult stem cells from endometrial biopsy(2). Based on this, Margherita Y. Turco, et al developed noninvasive methods to construct endometrial organoids from menstrual flow(4). Apart from adult stem cells, pluripotent stem cells were also induced to endometrial stromal fibroblasts and epithelium, and then cocultured to form organoids(5, 6), which offered vigorous proliferative capacity but lacked immune cells and other components of the microenvironment. There are also studies that add immune cells to endometrial organoids for co-culture(7). In efforts to mimic human endometrial cell types in vitro, the endometrial organoids progressed from epithelial organoids(8), to assemblies of epithelial and stromal cells(9, 10) and then to stem cell-laden 3D artificial endometrium(11, 12), which were solely closer but not completely identical to the endometrium. Additionally, pathological endometrial organoid models have been established for conditions like endometriosis, endometrial hyperplasia, Lynch syndrome, and endometrial cancer (13). These organoids simulated endometrial morphology, hormone responsiveness, and physiological and pathological processes in vitro (3, 8, 14–16), facilitating the study of physiological phenomena (14, 17, 18), pathogenic mechanisms(15) and drug screening(13).

Earlier version: Although various regulators have been suggested to be involved in WOI, we are far from understanding how embryo implantation occurs during the WOI due to ethical limitations and lack of in vitro receptive endometrial model. Transforming from proliferative to receptive endometrium involves dynamic changes like decidualization, epithelial-mesenchymal transition (EMT) and ciliated epithelial development (1, 18), which have not been reconstructed in any in vitro model to date. Here, we successfully established a receptive endometrial organoid system regulated by hormones and reflecting the in vivo WOI endometrial features, providing a platform for studying physiological and pathological endometrium and maternal-fetal interactions during the peri-implantation period.

Revised version: Although various regulators have been suggested to be involved in WOI, we are far from understanding how embryo implantation occurs during the WOI due to ethical limitations and fewer in vitro receptive endometrial model. Transforming from proliferative to receptive endometrium involves dynamic changes like decidualization, epithelial-mesenchymal transition (EMT) and ciliated epithelial development (1, 19), which have not been reconstructed in most in vitro model to date. Here, we successfully established a receptive endometrial organoid system regulated by hormones and reflecting the in vivo WOI endometrial features, providing a platform for studying physiological and pathological endometrium and maternal-fetal interactions.

Results

Endometrial organoids possess endometrial morphology and function

Earlier version: To establish endometrial organoids in vitro, prereceptive endometrium from reproductive-age women was dissociated into single cells or small cell masses, which were embedded with Matrigel in expansion medium (ExM). We devised a “window-of-implantation (WOI)” organoid by supplementing the culture system with prolactin (PRL), human chorionic gonadotropin (hCG), and human placental lactogen (hPL), alongside E2, MPA, and cAMP. In contrast, endometrial organoids maintained in ExM constituted the “control (CTRL)” group, while organoids treated with E2 for two days, followed by E2, MPA, and cAMP for an additional six days, induced the secretory phase as previously reported(7), and were designated the “secretory (SEC)” group (Fig. 1A). Theoretically, the WOI organoids are initially in the secretory phase, thus, possessing characteristics of the SEC organoid, more importantly, they should exhibit the traits of mid-secretory phase. The endometrial organoid, consisting of vesicle-like glands, fibrous stromal cells and other surrounding cells, developed into a 3D structure with the support of Matrigel (Fig. 1B∼1C, Video S1∼S4). As the organoids grew and differentiated, the endometrial glands enlarged, epithelial cells adopted a paving stone arrangement, and stromal cells formed an extensive network (Fig. 1C). The glandular epithelium marker E-cadherin, stromal cell marker vimentin, and endometrial gland marker FOXA2 were expressed in the cultured endometrial organoid, which resembled the endometrium morphologically (Fig. 1D). Moreover, the endometrial organoids exhibited substantial expression of the proliferation marker Ki67, while the apoptosis marker cleaved caspase-3 was undetectable, signifying the organoids’ robust proliferative capacity (Fig. 1E).

Revised version: To establish endometrial organoids in vitro, pre-receptive endometrium from reproductive-age women was dissociated into single cells or small cell masses. These cells then self-assembled into organoids induced by various small molecules, such as Noggin, EGF, FGF2, WNT-3A and R-Spondin1 in expansion medium (ExM) (Fig.1A). The endometrial organoids consist of vesicle-like glands, fibrous stromal cells, and other surrounding cells, and they were capable of being passaged (Fig. 1B∼1C, Fig. S1A∼S1B, Video S1∼S4). As the organoids grew and differentiated, the endometrial glands enlarged, epithelial cells adopted a paving stone arrangement, and stromal cells formed an extensive network (Fig. 1C). The epithelium marker E-cadherin, stromal cell marker vimentin, and endometrial gland marker FOXA2 were expressed in the cultured endometrial organoid, which resembled the endometrium morphologically (Fig. 1D). Moreover, the endometrial organoids exhibited substantial expression of the proliferation marker Ki67, while the apoptosis marker cleaved caspase-3 was undetectable, signifying the organoids’ robust proliferative capacity (Fig. 1E).

Endometrial organoids possess endometrial morphology and function

(A) Human endometrial organoids constructed from adult stem cells were treated with expansion medium (ExM) (CTRL) or subjected to hormonal stimulation. Timeline of endometrial organoid cultured by ExM (CTRL), ovarian steroid hormones simulating secretory phase (SEC), ovarian steroid hormones combining PRL and placental hormones to mimic the window of implantation (WOI).

(B) Endometrial organoids in the CTRL, SEC and WOI groups displayed similar growth patterns during the culture period. Scale bar = 200 μm.

Earlier version: (C) Screenshot of video S1 showing endometrial glands gradually developing into a vesicular shape, and the surrounding stromal cells arranging in a fibrous pattern in the endometrial organoids (100x, up-left). Screenshot of video S2 displaying the stromal cells growing in fibrous pattern and forming an extensive network in the endometrial organoids (200x, up-right). (The yellow arrows indicate stromal cells and the white arrows indicate endometrial glands). The epithelial cells arrange like paving stones (middle and down, left). Stromal cells formed an extensive network (middle and down, right). The arrowhead indicates stromal cells. Scale bar = 100 μm (middle), Scale bar = 50 μm (down).

Revised version: (C) Screenshot of video S1 showing endometrial glands gradually developing into a vesicular shape, and the surrounding stromal cells arranging in a fibrous pattern in the CTRL endometrial organoids (100x, up-left). Screenshot of video S2 displaying the stromal cells growing in fibrous pattern and forming an extensive network in the CTRL endometrial organoids (200x, up-right). (The yellow arrows indicate stromal cells and the white arrows indicate endometrial glands). The epithelial cells arrange like paving stones (middle and down, left). Stromal cells formed an extensive network (middle and down, right). The arrowhead indicates stromal cells. Scale bar = 100 μm (middle), Scale bar = 50 μm (down).

Earlier version: (D) Validation of epithelial, stromal cell and endometrial gland markers (E-cadherin, vimentin and FOXA2, respectively) with immunofluorescence (IF) in the endometrium in vivo and endometrial organoids in vitro. Nuclei were counterstained with DAPI. Scale bar = 40 μm. of endometrium in vivo and endometrial organoid in vitro. Scale bar = 20 μm.

Revised version: (D) Validation of epithelial, stromal cell and endometrial gland markers (E-cadherin, vimentin and FOXA2, respectively) with immunofluorescence (IF) in the endometrium in vivo and CTRL endometrial organoids in vitro. Nuclei were counterstained with DAPI. Scale bar = 40 μm. Periodic acid-schiff staining (PAS) of endometrium in vivo and endometrial organoid in vitro. Scale bar = 20 μm.

Earlier version: (E) IF analysis of proliferation and apoptosis indicated by Ki67 and cleaved caspase-3, respectively. Nuclei were counterstained with DAPI. Scale bar = 40 μm.

Revised version: (E) IF analysis of proliferation and apoptosis indicated by Ki67 and cleaved caspase-3 in the CTRL organoids, respectively. Nuclei were counterstained with DAPI. Scale bar = 40 μm.

Earlier version: Functionally, endometrial organoids successfully secreted glycogen into the lumen, analogous to in vivo endometrial activity, thereby providing nourishment for embryo implantation (Fig. 1D). The CTRL organoids exhibit minimal glycogen secretion, whereas both SEC and WOI organoids display substantial glycogen secretion into the glandular lumen (Fig.S1A). Furthermore, the endometrium undergoes cyclical alterations under regulation by estrogen (E2) and progesterone in vivo. Following hormonal stimulation, WOI organoids exhibited slower growth than SEC and CTRL organoids, while CTRL organoids maintained robust proliferative activity (Fig. 1B). In our organoids, after treatment with E2 for two days and subsequently with medroxyprogesterone acetate (MPA) and cAMP for fourteen days, the organoids demonstrated significantly increased expression of progesterone receptor (PR A/B) and a slight elevation in estrogen receptor α (ERα) relative to the CTRL group (Fig. S1B), suggesting the sensitivity to hormone treatment. These results demonstrate that our organoids exhibited comparable morphology, secretory function, and hormone responsiveness to the in vivo endometrium.

Revised version: Functionally, endometrial organoids effectively secreted glycogen into the lumen, mirroring the in vivo endometrial activity, thereby providing nourishment for embryo implantation (Fig. 1D). Moreover, after treatment with estrogen (E2) for two days followed by administration of medroxyprogesterone acetate (MPA) and cAMP for fourteen days, the organoids exhibited a notable upregulation of progesterone receptor (PR A/B) expression and a modest increase in estrogen receptor α (ERα) compared to the control group (Fig. S1C∼S1D). It reflected a comparable hormonal responsiveness in the in vivo endometrium.

Added in the revised version: To identify hormonal regimens that induce implantation window, pregnancy-related hormones were supplemented into the culture system following the two-day E2. E2 and MPA promote the transition of endometrial organoids into the secretory phase, while pregnancy hormones can promote the further differentiation. Prolactin (PRL) promotes immune regulation and angiogenesis during implantation(2, 20). Human chorionic gonadotropin (hCG) improves endometrial thickness and receptivity (21, 22). Human placental lactogen (hPL) promotes the development and function of endometrial glands(23). Hormone dosage was primarily based on peri-pregnant maternal body or localized endometrium levels (2). A comparison of multiple groups revealed similar counts, area, and average intensity of organoids over time (Fig.S1E). However, only the final recipe (i.e. a combination of E2, MPA, cAMP, PRL, hCG and hPL) exhibited an endometrial receptivity-related gene expression profile, which highly expressed genes positively correlated with endometrial receptivity, and lowly expressed genes negatively correlated with receptivity, compared to the other hormone formulations (Fig.S1F). The organoids induced with this scheme were defined as WOI organoids (Fig. 1A). In contrast, endometrial organoids maintained in ExM constituted the “control (CTRL)” group, whereas organoids treated with E2 for two days, followed by E2, MPA, and cAMP for an additional six days, were induced to the secretory phase as previously reported(8), and designated as the “secretory (SEC)” group (Fig. 1A). Theoretically, the WOI organoids are initially in the secretory phase, thus sharing characteristics with the SEC organoids. More importantly, they are expected to display traits typical of the mid-secretory phase. There was no significant difference in the morphology of organoids among the three groups (Fig.1B). The glycogen secretion of CTRL, SEC and WOI organoids showed an increasing trend (Fig.S1G).

Developing receptive endometrial organoids in vitro mimicking the implantation-window endometrium

Earlier version: Single-cell transcriptomics analysis with reference to CellMarker, PanglaoDB, Human Cell Atlas, Human Cell Landscape, and scRNASeqDB, and previous endometrium related studies(1, 7, 9, 19) revealed the presence of epithelium, stromal cells and immune cells in WOI organoids (Fig. 2A, Fig. S3B∼S3C). With respect to glandular and luminal epithelium, secretory epithelium, LGR5 epithelium, ciliated epithelium and EMT-derived stromal cells, the WOI organoids closely resemble the mid-secretory endometrium (Fig.2A, Fig.S3C∼S3D) according to comparisons of the scRNA-seq data of our organoids and mid-secretory endometrium described by Stephen R. Quake in 2020 (1). We have captured the morphology of immune and stromal cells using 3D clearing staining and light sheet microscopy imaging, with vimentin marking stromal cells, CD44 designating immune cells and FOXA2 identifying glands (Fig.2B∼2C). Thus, we sought to confirm the presence of immune cells and investigate the proportion of T cells and macrophages in the endometrial organoids using flow cytometry. The gating strategy for T cells and macrophages was shown in Fig. 2D. White blood cells (WBC) were determined as CD45+ cells, while T cells and macrophages were defined as CD45+CD3+ cells and CD45+CD68+CD11b+ cells, respectively.

Revised version: Single-cell transcriptomics analysis identified the presence of epithelium, stromal cells, and immune cells in WOI organoids with reference to CellMarker, PanglaoDB, Human Cell Atlas, Human Cell Landscape, and scRNASeqDB, and previous endometrium related studies(1, 8, 10, 24) (Fig. 2A, Fig. S2A∼S2B). The WOI organoids exhibited similarities to the mid-secretory endometrium in terms of glandular and luminal epithelium, secretory epithelium, LGR5 epithelium, ciliated epithelium, and EMT-derived stromal cells, as evidenced by comparisons of scRNA-seq data from our organoids and the mid-secretory endometrium as described by Stephen R. Quake in 2020 (1) (Fig.2A, Fig.S2C∼S2D). The morphology of immune and stromal cells was analyzed through 3D clearing staining and light sheet microscopy imaging, with vimentin labeling stromal cells, CD45 and CD44 indicating immune cells, and FOXA2 identifying glands (Fig.2B∼2C). Furtherly, we confirmed the presence of immune cells and investigated the proportion of T cells and macrophages using flow cytometry. White blood cells (WBC) were identified as CD45+ cells, with T cells and macrophages characterized as CD45+CD3+ cells and CD45+CD68+CD11b+ cells, respectively (Fig. 2D). The proportion of WBC in organoids was approximately 3%∼4% (Fig.2D), among which macrophages were less than 1% and T cells less than 2% (Fig. S2E). The percentage of WBCs in WOI organoids was lower than that in CTRL organoids, among which T cells decreased in the WOI organoids (Fig. S2F).

Developing receptive endometrial organoids in vitro mimicking the implantation-window endometrium

(A) T-SNE plot of scRNA-seq data from three individual endometrial organoids of the CTRL, SEC and WOI groups (left). T-SNE plot of combined scRNA-seq data from the three kinds of organoids and mid-secretory endometrium (right).

Earlier version: (B) Exhibition of stromal cell marked by vimentin through combination of organoid clearing, IF and light sheet microscopy. Nuclei were counterstained with DAPI. The arrowhead indicates stromal cells. Scale bar = 40 μm (left), Scale bar = 30 μm (right).

Revised version: (B) Exhibition of stromal cell marked by vimentin of CTRL organoid through whole-mount clearing, immunostaining and light sheet microscopy imaging. Nuclei were counterstained with DAPI. The arrowhead indicates stromal cells. Scale bar = 40 μm (left), Scale bar = 30 μm (right).

Earlier version: (C) Exhibition of immune cell marked by CD44 and endometrial gland marked by FOXA2 through combination of organoid clearing, IF and light sheet microscopy. Nuclei were counterstained with DAPI. The arrowhead indicates immune cells. Scale bar = 30 μm (left), Scale bar = 10 μm (right).

Revised version: (C) Exhibition of immune cell marked by CD45 and CD44, and endometrial gland marked by FOXA2 of CTRL organoid through whole-mount clearing, immunostaining and light sheet microscopy imaging. Nuclei were counterstained with DAPI. The arrowhead indicates immune cells. Scale bar = 50 μm (left), Scale bar = 10 μm (right).

Earlier version: (D) Flow cytometric analysis of T cells and macrophages in the endometrial organoid. Gating strategy used for determining white blood cells (CD45+ cells), T cells (CD45+CD3+ cells) and macrophages (CD45+CD68+CD11b+ cells).

Revised version: (D) Flow cytometric analysis of T cells and macrophages in the CTRL endometrial organoid. Gating strategy used for determining white blood cells (WBC) (CD45+ cells), T cells (CD45+CD3+ cells) and macrophages (CD45+CD68+CD11b+ cells).

(E) Heatmap and bubble diagram illustrating highly expressed genes as well as GO functions enriched in both organoids during the WOI and mid secretory endometrial tissue in terms of SOX9+LGR5+ epithelium, stem-derived epithelium, secretory epithelium, proliferative epithelium, unciliated epithelium, stromal cells and EMT-derived stromal cells. The color of heatmap represents log-transformed fold change of gene expression.

(F) Heatmaps showing differentially expressed TFs of endometrial organoids and endometrium in the secretory epithelium (left) and EMT-derived stromal cells (right). The color represents log-transformed fold change of gene expression.

(G) Endometrial receptivity evaluation of endometrium and their derived WOI organoids through ERT. Asterisks indicate individual samples.

Added in the revised version: (H) Electron micrograph of the CTRL (top), SEC (middle) and WOI (bottom) endometrial organoid showing pinopodes (P), glycogen granule (asterisk), microvilli (white arrows) and cilia (orange arrows). Scale bar = 1 μm. Quantitative comparison of pinopodes, glycogen, microvilli, and cilia in the CTRL, SEC and WOI organoids. *P ≤ 0.05, **P ≤ 0.005, ***P ≤ 0.0005, ****P ≤ 0.0001.

Removed in the revised version: Given that the mid-secretory phase is when the implantation window opens and significant transcriptional changes occur, we focused on comparing gene expression and transcription regulation during this phase (Fig.2E∼2F), referring to Stephen R. Quake 2020 for mid-secretory phase scRNA-seq data(1). Pathways including mitochondrial energy metabolism, cell adhesion, cytoskeleton regulation, and epithelial cell growth were all upregulated in both the WOI organoid and mid secretory endometrium compared with that in the CTRL and SEC organoids (Fig. 2E).

Earlier version: Analysis of key transcription factors (TFs) involved in implantation revealed similarities between WOI organoids and mid secretory epithelium when compared to the CTRL and SEC organoids, including hypoxia response (such as previously reported HIF1A(20)), embryo implantation (such as FBLN1(21)), lipid metabolism (such as VMP1(22)), cell migration, and cell junction (such as TJP1(23)) (Fig. 2F). Moreover, markers of endometrial decidualization (such as S100A10(24)), EMT (such as FAT1(25) and FOXF2(26)), and receptivity (such as NEAT1(27), SERPINB9(28), SOX17(29) and SOX4(30)) were expressed at comparable level in EMT-derived stromal cells in both groups (Fig. 2F). We then conducted the endometrial receptivity test (ERT) to assess the receptive state of the endometrium(31) and endometrial organoids (Fig. 2G). Hormone-treated organoids derived from pre-receptive endometrium entered the receptivity phase, suggesting the successful formation of the implantation window in vitro. Collectively, these results demonstrate that the developed WOI organoids closely resemble the in vivo endometrium during the implantation window in terms of cell composition, function, transcriptional regulation, and endometrial receptivity.

Revised version: Our study concentrated on the comparison of gene expression and transcriptional regulation during the mid-secretory phase, as this phase marks the implantation window and involves substantial transcriptional alterations (Fig.2E∼2F). We referred to the scRNA-seq data of the mid-secretory phase from Stephen R. Quake 2020 (1). Pathways related to mitochondrial energy metabolism, cell adhesion, cytoskeleton regulation, and epithelial cell growth exhibited upregulation in both the WOI organoid and mid-secretory endometrium in comparison to the CTRL and SEC organoids (Fig. 2E). The crucial transcription factors (TFs) found in the secretory epithelium and EMT-derived stromal cells, which are implicated in implantation, revealed similarities between WOI organoids and mid-secretory endometrium. Specifically, the secretory epithelium exhibited comparable TFs related to hypoxia response (such as previously reported HIF1A(25)), embryo implantation (such as FBLN1(26)), lipid metabolism (such as VMP1(27)), cell migration, and cell junction (such as TJP1(28)) (Fig. 2F). Similarly, EMT-derived stromal cells also expressed similar TFs involved in endometrial decidualization (such as S100A10(29)), EMT (such as FAT1(30) and FOXF2(31)), and receptivity (such as NEAT1(32), SERPINB9(33), SOX17(34) and SOX4(35)) (Fig. 2F).

Added in the revised version: We then conducted the endometrial receptivity test (ERT) (36) and explored microstructure to further assess the receptive state of WOI organoids (Fig. 2G). ERT is a kind of gene analysis-based method for detecting endometrial receptivity, which combines high-throughput sequencing and machine learning to analyze the expression of endometrial receptivity-related genes (36). It is currently used in clinical practice to determine endometrial receptivity and guide personalized embryo transfer. The WOI organoids derived from pre-receptive endometrium were observed to transit into the receptivity phase (Fig. 2G). Besides, the WOI organoids possessed various characteristic microstructure of implantation window, including elongated microvilli and increased glycogen, pinopodes, and especially cilia (Fig.2H).

Added in the revised version: Collectively, these results demonstrated that the WOI organoids closely resemble the in vivo endometrium during the implantation window in terms of cell composition, function, transcriptional regulation, and endometrial receptivity.

Receptive endometrial organoids recapitulate WOI-associated biological characteristics

Earlier version: We analyzed the transcriptome and proteome profiles of WOI organoids compared to those of CTRL or SEC organoids to assess WOI associated biological characteristics (Fig. 3A, 3E, Fig. S2A∼S2D). WOI organoids exhibited a strong hormone response, as demonstrated by PGR at the transcriptome level (Fig. 3B). Combined analysis of the transcriptome and proteome of the WOI organoids demonstrated 179 upregulated genes/proteins compared with those in the CTRL group (Fig. S2E), and most of these genes/proteins are involved in the estrogen response (Fig. S2F). Furthermore, the progesterone response level tested by PRA/B immunostaining was highest in the WOI group, meanwhile estrogen responsive protein OLFM4 was also upregulated (Fig. 3C). FOXO1 is a critical marker of endometrial receptivity depending on PGR signaling. We found WOI organoids showed a significantly elevated FOXO1 compared with the CTRL organoids (Fig.3D, Video S5), indicating that progesterone signaling was involved in WOI organoid establishment. The above results demonstrated that the WOI organoids responded strongly to estrogen and progesterone.

Revised version: We analyzed the transcriptome and proteome profiles of WOI organoids compared to those of CTRL or SEC organoids to assess WOI associated biological characteristics (Fig. 3A, 3E, Fig. S3A∼S3D). WOI organoids exhibited a robust hormone response, as demonstrated by the upregulation of PGR at the transcriptome level (Fig. 3B). Integrated analysis of the transcriptome and proteome data identified 179 upregulated genes/proteins in WOI organoids compared to the CTRL group (Fig. S3E), with most of these genes/proteins implicated in the estrogen response (Fig. S3F). Furthermore, the analysis of progesterone response levels using PRA/B immunostaining indicated the highest levels in the WOI group, accompanied by an upregulation of the estrogen responsive protein OLFM4 (Fig. 3C). FOXO1, a crucial marker of endometrial receptivity reliant on PGR signaling, exhibited significantly elevated expression in WOI organoids compared to CTRL organoids (Fig.3D, Video S5), suggesting the involvement of progesterone signaling in the establishment of WOI organoids. These findings collectively demonstrated a strong response of WOI organoids to estrogen and progesterone.

Receptive endometrial organoid exhibits the implantation window at the transcription and protein levels

(A) Principal component analysis (PCA) plot computed with differentially expressed genes in the bulk transcriptome of endometrial organoids belonging to the CTRL, SEC and WOI groups.

(B) Heatmap showing that enrichment of differentially expressed genes for the terms of extracellular matrix remodeling, hormone response, negative regulation of cell proliferation, monocarboxylic acid metabolism, lipid metabolism, cell adhesion, negative regulation of cell differentiation and RHO GTPase signaling. The color represents log-transformed fold change of gene expression.

(C) Responsiveness to progesterone and estrogen was evaluated by IF to PRA/B and OLFM4 with IF respectively. Scale bar = 40 μm, **P ≤ 0.005.

(D) Exhibition of implantation marker (FOXO1) and endometrial gland marker (FOXA2) through combination of organoid clearing, IF and light sheet microscopy. Nuclei were counterstained with DAPI. **P ≤ 0.005.

(E) PCA plot computed with differentially expressed proteins in the microproteomics of endometrial organoids belonging to the CTRL, SEC and WOI groups.

(F) Dot bar diagram exhibiting protein expression levels related to cilia in the CTRL and WOI groups (left). Dot bar diagram displaying tubulin, integrin and PI3K-AKT pathway-related protein expression levels in the SEC and WOI groups (right).

(G) IF analysis of cilia assembly marked by acetyl-α-tubulin. Nuclei were counterstained with DAPI. The arrowhead indicates cilia. Scale bar = 50 μm. *P ≤ 0.05.

Removed in the revised version: (H) Electron micrograph of the CTRL (top) and WOI (bottom) endometrial organoid showing pinopodes (P), glycogen granule (asterisk), microvilli (white arrows) and cilia (orange arrows). Scale bar = 1 μm.

Earlier version: The growth and development, assembly and movement of cilia, a characteristic endometrial structure, are dynamically regulated by the hormones during the menstrual cycle, which are essential for the establishment of the endometrial implantation window and embryo implantation. Cilia assembly(marked by acetyl-α-tubulin(13)) and motile cilia (marked by TPPP3 and ODF2(32)) appeared dominant in the WOI organoids (Fig. 3F∼3G, Fig.S2I). Moreover, the expression of tubulin, integrin and components of the PI3K-Akt pathway, which function to generate pinopodes and respond to invasion, was also upregulated in the WOI organoids (Fig. 3F, Fig.S2J). Electron microscopy confirmed the existence of pinopodes, glycogen particles, microvilli and cilia in the endometrial organoids (Fig. 3H). The pinopodes of WOI organoids arranged more densely on the luminal side of epithelium than that of CTRL organoids, which contributes the embryo adhesion. The endometrial epithelium is composed of ciliated cells and non-ciliated cells with microvilli. The microvilli of WOI organoids were densely packed, while cilia were only observed in the WOI organoids. In addition, the WOI organoids had much more glycogen particles than the CTRL group. These structures marked the establishment of endometrial receptivity and benefited for the embryo adhesion to the endometrium.

Revised version: The growth and development, assembly and movement of cilia, a characteristic endometrial structure, are essential for the establishment of endometrial implantation window and embryo implantation. Cilia assembly(marked by acetyl-α-tubulin(14)) and motile cilia (marked by TPPP3 and ODF2(37)) were prominently observed in the WOI organoids (Fig. 3F∼3G, Fig. S3I). Moreover, the expression of tubulin, integrin, and components of the PI3K-Akt pathway was also upregulated in the WOI organoids, which functioned to generate pinopodes, and respond to invasion (Fig. 3F, Fig. S3J). Under the electron microscopy, the cilia were most observed in the WOI organoids, while the pinopodes of WOI organoids exhibited the densest arrangement on the luminal side of epithelium among these three groups (Fig. 2H). These structures marked the establishment of endometrial receptivity and benefited for the embryo adhesion to the endometrium.

Earlier version: Other characteristics of the implantation window involved decreased cell proliferation (marked by CDH5 and MEIS1), increased cell differentiation (marked by FGFBP1 and TFF1), extracellular matrix (ECM) remodeling (marked by PDGFB and IGFBP4(33)) and reduced cell adhesion (marked by EFNA5, LAMA3 and RHO GTPase) (Fig. 3B, Fig. S2G). Besides, energy metabolism (represented by SLC25A1(34)) and hypoxia response (represented by HIF1 α (20)) were also characteristically enhanced in WOI organoids (Fig. 3B, Fig. S2F∼S2H), which maintained the receptive state of endometrial organoids.

Revised version: Other characteristics of the implantation window involved decreased cell proliferation (marked by CDH5 and MEIS1), increased cell differentiation (marked by FGFBP1 and TFF1), extracellular matrix (ECM) remodeling (marked by PDGFB and IGFBP4(38)) and reduced cell adhesion (marked by EFNA5, LAMA3 and RHO GTPase) (Fig. 3B, Fig. S3G). Besides, energy metabolism (represented by SLC25A1(39)) and hypoxia response (represented by HIF1α(25)) were also characteristically enhanced in WOI organoids (Fig. 3B, Fig. S3F∼S3H), which maintained the receptive state of endometrial organoids.

Added in the revised version: Overall, although the WOI organoids possessed the characteristics of implantation window, the differences in transcriptome and proteome between SEC and WOI organoids at the organoid level are not significant. This is understandable as WOI organoids are further induced towards the implantation window based on the secretory phase (i.e. SEC organoids), which prompted us to continue exploring at the single-cell level.

Structural cells construct WOI with functionally dynamic changes

Earlier version: Based on the overall WOI associated biological characteristics at the organoid level, we further examined the functionally dynamic changes in structural cells during the WOI. Although the differences between the WOI and SEC organoids were slight at the organoid level, the WOI organoids displayed more thoroughly differentiation to the mid-secretory endometrium than the SEC organoids at the single-cell level.

Revised version: Based on the WOI associated biological characteristics at the organoid level, we further examined the functionally dynamic changes in structural cells during the implantation window. The WOI organoids displayed more thoroughly differentiation to the mid-secretory endometrium than the SEC organoids at the single-cell level.

Earlier version: In terms of energy metabolism, the WOI organoids exhibited upregulation of monocarboxylic acid and lipid metabolism and hypoxia response. Likewise, the secretory epithelium, critical for the implantation window, accumulated more in the SEC and WOI groups, contributing to cellular metabolic processes and HIF-1 signaling pathway response to hypoxia at the single-cell level (Fig. 4A, Fig. S3E). WOI organoid secretory epithelium exhibited enhanced peptide metabolism and mitochondrial energy metabolism compared to the SEC group, supporting endometrial decidualization and embryo implantation (Fig. 4B). Moreover, increased secretory epithelium and decreased proliferative epithelium in the SEC and WOI organoids suggested the transformation from proliferation to secretory phase, which is an obvious feature of the implantation window (Fig. 4A). Proliferative epithelium differentiates into secretory epithelium under the regulation of branch nodal genes between state 5 and state 6, such as the GAS5, RPL and RPS families, as revealed by single-cell trajectory (Fig. 4C, Fig. S4A∼B, Fig. S4G). The WOI organoids showed more thoroughly differentiation from proliferative to secretory epithelium in relation to the SEC organoids (Fig. 4C, Fig. S4A). Overall, the WOI organoids indicated that proliferative epithelium transformed to secretory epithelium with increased energy metabolism and hypoxia response during the implantation window.

Revised version: In terms of energy metabolism, the WOI organoids exhibited upregulation of monocarboxylic acid and lipid metabolism, and hypoxia response (Fig.3B, Fig.S3F). Likewise, the secretory epithelium, critical for the implantation window, contributed to cellular metabolic processes and HIF-1 signaling pathway response to hypoxia at the single-cell level (Fig. 4A, Fig. S2G). The secretory epithelium of WOI organoids exhibited enhanced peptide metabolism and mitochondrial energy metabolism compared to the SEC group, supporting endometrial decidualization and embryo implantation (Fig. 4B). Moreover, increased secretory epithelium and decreased proliferative epithelium in the SEC and WOI organoids suggested the transformation from proliferation to secretory phase (Fig. 4A). Proliferative epithelium differentiates into secretory epithelium under the regulation of branch nodal genes between state 5 and state 6, such as the GAS5, RPL and RPS families, as revealed by single-cell trajectory (Fig. 4C, Fig. S4A∼B, Fig. S4G). The WOI organoids showed more thoroughly differentiation from proliferative to secretory epithelium in relation to the SEC organoids (Fig. 4C, Fig. S4A). Overall, the WOI organoids indicated more comprehensive differentiation to the secretory epithelium with increased energy metabolism and hypoxia response during the implantation window.

Receptive endometrial organoid exhibits the implantation window at single cell level

(A) Bar graph exhibiting various percentages of each cell type in the three groups.

(B) Gene set enrichment analysis (GSEA) between the SEC and WOI groups for the secretory epithelium.

(C-E) Pseudotime trajectory showing the transformation between proliferative and secretory epithelium (C), ciliated and unciliated epithelium (D), proliferative epithelium and EMT-derived stromal cells (E) in the CTRL, SEC and WOI groups. Arrows indicate the direction of the pseudotime trajectory.

(F-G) Dot plots demonstrating the Cellphone DB analysis of relevant receptors and ligands of ciliated epithelium (F) and EMT-derived stromal cells (G) with other cell types. The size of the dot represents the level of significance. The color of the dot indicates the mean of the average expression level of interacting molecule 1 in ciliated epithelium or EMT-derived stromal cells and molecule 2 in other cell types.

Earlier version: Given that the WOI organoids exhibited increased cilia assembly and motile cilia, the transformation and function of the ciliated epithelium gained attention. Single-cell transcriptome analysis revealed hormone treatment increased ciliated epithelium and decreased unciliated epithelium in SEC and WOI groups by altering the expression of key genes (Fig. 4A). These crucial genes, including KRT19, MALAT1, RPL21 and RPS2, expressed divaricately at the fork of differentiation from ciliated cells toward ciliated or unciliated cells (Fig. 4D, Fig. S4C∼D, Fig. S4H). Ciliated epithelium was enriched in protein binding, cilium organization, and assembly genes, while unciliated epithelium was enriched in actin cytoskeleton and translation genes (Fig. S3E). The WOI organoids’ ciliated epithelium regulated vasculature development and displayed higher transcriptional activity than the SEC group (Fig. S3F). Ciliated epithelium interacts with immune cells and secretory epithelium, modulating angiogenesis via NRP1-VEGFA and NRP2-VEGFA(35) in all groups, with NRP2 highly expressed in WOI organoids (Fig. 4F, Fig. S2K). The WOI ciliated epithelium showed enhanced invasion ability via CD74-COPA(36), CD74-APP(36) and ROR2-WNT5A(37) (Fig. 4F). Thus, WOI organoids revealed accumulated ciliated epithelium’s role in preparing the implantation window.

Revised version: Given that the WOI organoids exhibited increased cilia assembly and motile cilia, the transformation and function of the ciliated epithelium gained attention. Single-cell transcriptome analysis revealed hormone treatment increased ciliated epithelium and decreased unciliated epithelium in SEC and WOI groups by altering the expression of key genes, such as KRT19, MALAT1, RPL21 and RPS2 (Fig. 4A, 4D, Fig. S4C∼D, Fig. S4H). Ciliated epithelium functioned in protein binding, cilium organization and assembly, while unciliated epithelium acted on actin cytoskeleton and translation (Fig. S2G). The WOI organoids’ ciliated epithelium regulated vasculature development and displayed higher transcriptional activity than the SEC group (Fig. S2H). The implementation of these functions depended on its interaction with other cells. CellPhoneDB is a useful tool to investigate ligand-receptor interactions between the cells. Ciliated epithelium interacted with immune cells and secretory epithelium, modulating angiogenesis via NRP1-VEGFA and NRP2-VEGFA(40), with NRP2 highly expressed in the WOI organoids (Fig. 4F, Fig. S3K). Besides, the ciliated epithelium of WOI organoids showed enhanced invasion ability via CD74-COPA(41), and ROR2-WNT5A(42), which was validated by proximity ligation assay (PLA) (Fig. 4F, Fig. S2J). Thus, WOI organoids revealed accumulated ciliated epithelium’s role in preparing the implantation window.

Earlier version: The WOI organoids displayed upregulated cell differentiation not only at organoid level but also at cellular level, which is represented by EMT. EMT is a common and crucial biological event in the endometrium during the implantation window(38). Synthetic analysis of the transcriptome and proteome revealed increased EMT occurrence in WOI organoids (Fig. S2F). EMT-derived stromal cells, exhibiting gene expression patterns characteristic of both epithelial and stromal cells, are the indication of EMT and more abundant in the SEC and WOI groups (Fig. 4A). Pseudotime trajectory indicated proliferative epithelium transformed into EMT-derived stromal cells, regulated by key genes such as DOC2B, FXYD3, and LPCAT3 (Fig. 4E, Fig. S4E∼F, Fig. S4I). The WOI organoids exhibited more thoroughly differentiation to EMT-derived stromal cells than the SEC organoids (Fig. 4E). EMT-derived stromal cells showed alterations in protein binding, cell cycle, organelle organization, and reproduction, with more enhanced lamellipodium-mediated cell migration, cell junction and cytoskeleton regulation in the WOI organoids than the SEC organoids (Fig. S3E∼S3F). EMT-derived stromal cells and epithelium cooperated in transitioning from proliferation phase to implantation window (Fig.4G). We found that NRP1 and SLC7A1 were highly expressed by EMT-derived stromal cells, and their receptors (SEMA3A and CSF1) were more upregulated in the epithelium of WOI group than SEC group (Fig.4G). NRP1-SEMA3A has been reported to promote vascularization and responds to hypoxia(39). SLC7A1(40) is the key arginine transporter, and CSF1(41) is a kind of pro-survival cytokines, both of which support receptivity establishment, embryo implantation and development. However, angiogenesis-related and TNF-related pathways were downregulated in WOI organoids (Fig.4G). The functions of NRP1-VEGFB and ADRB2-VEGFB in angiogenesis are not yet clear and require further exploration. Elevated TNF has been reported connected to stromal cell demise and pregnancy failure(42). Compared with common stromal cells, EMT-derived stromal cells communicate slightly differently with epithelial or immune cells. CD44 and CD46, which are involved in cell adhesion(43) and immunoregulation(44), are highly expressed in stromal cells and EMT-derived stromal cells, respectively, and bind separately with SPP1 and JAG1 in epithelial and stromal cells (Fig. 4G, Fig. S3G). In general, WOI organoids demonstrated EMT’s role in mediating endometrial transformation towards the implantation window.

Revised version: The WOI organoids displayed upregulated cell differentiation not only at organoid level but also at cellular level, which is represented by EMT. EMT is a common and crucial biological event in the endometrium during the implantation window(43). During the EMT process, epithelial cells lose their epithelial characteristics while gaining migratory and invasive properties of fibroblasts. Synthetic analysis of the transcriptome and proteome revealed increased EMT in WOI organoids (Fig. S3F). EMT-derived stromal cells, exhibiting gene expression patterns typical of epithelial and stromal cells, as well as EMT, are more abundant in the SEC and WOI groups, and act on protein binding, cell cycle, organelle organization, and reproduction (Fig. 4A, Fig. S2G). They performed enhanced lamellipodium-mediated cell migration, cell junction and cytoskeleton regulation in the WOI organoids compared to the SEC organoids (Fig. S2H). The WOI organoids exhibited more thoroughly differentiation from proliferative epithelium to EMT-derived stromal cells than the SEC organoids, which was regulated by key genes such as DOC2B, FXYD3, and LPCAT3 (Fig. 4E, Fig. S4E∼F, Fig. S4I). EMT-derived stromal cells and epithelium cooperated in transitioning from proliferation phase to implantation window (Fig.4G). We found that NRP1 and SLC7A1 were highly expressed by EMT-derived stromal cells, and their receptors (SEMA3A and CSF1) were more upregulated in the epithelium of WOI group than SEC group (Fig.4G, Fig. S2J). NRP1-SEMA3A has been reported to promote vascularization and responds to hypoxia(44). SLC7A1(45) and CSF1(46) both support receptivity establishment, embryo implantation and development. However, angiogenesis-related (NRP1-VEGFB and ADRB2-VEGFB), and TNF-related pathways were downregulated in WOI organoids (Fig.4G), which may link with pregnancy establishment and warrant further exploration (47). Compared with common stromal cells, EMT-derived stromal cells communicate slightly differently with epithelial or immune cells. CD44 and CD46, known for their roles in cell adhesion (48) and immunoregulation(49), are highly expressed in stromal cells and EMT-derived stromal cells, respectively, and bind separately with SPP1 and JAG1 in epithelial and stromal cells (Fig. 4G, Fig. S2I∼S2J). In general, WOI organoids demonstrated EMT’s role in mediating endometrial transformation towards the implantation window.

Discussion

In our study, we constructed the WOI endometrial organoids, and observed the remarkable resemblance in structure and function to the in vivo endometrium. The organoids consist of three primary types of endometrial cells, specifically epithelial, stromal, and immune cells, with epithelial cells assembling into glands and surrounded by immune cells and stromal cells. This similarity in cellular composition and tissue architecture to the endometrium lays the foundation for simulating receptive endometrium.

Earlier version: Although previous studies treated endometrial organoid with E2, P4, and cAMP to induce a transition to the secretory phase in the endometrium, the methodology to accurately simulate the implantation window or mid-secretory endometrium remains elusive and warrants further investigation. Previous research suggested that placental signals could further promote the differentiation of endometrial organoids (2), and PRL, hCG and HPL have been implicated in processes such as decidualization, implantation, immunoregulation, and angiogenesis (45–47). Detailly, PRL, synthesized by the adenohypophysis, endometrium, and myometrium, plays a vital role in implantation, immunoregulation, and angiogenesis, with the secretory endometrium producing PRL in response to MPA and E2, leading to ciliated cell formation and stromal cell decidualization (2, 45). HCG, secreted by trophoblasts in early pregnancy, influences decidual cells (46) and improves endometrial thickness and receptivity(48). The infusion of hCG increases the levels of key factors associated with endometrial receptivity, such as endocytosis proteins, hypoxia-inducible factor 1 (HIF1), chemokines, and glycodelin (49). HPL contributes to the development and function of uterine glands(47). Consequently, the additional supplementation of PRL, hCG and HPL in our system augmented hormone responsiveness and receptivity, while promoting cell differentiation, ultimately yielding a model more representative of the receptive endometrium.

Revised version: Although previous studies treated endometrial organoid with E2, P4, and cAMP to induce a transition to the secretory phase, the methodology to accurately simulate the implantation window or mid-secretory endometrium remains elusive and warrants further investigation. Previous research suggested that placental signals could further promote the differentiation of endometrial organoids (2), and PRL, hCG and HPL have been implicated in processes such as decidualization, implantation, immunoregulation, and angiogenesis (20, 21, 23). Detailly, PRL, synthesized by the adenohypophysis, endometrium, and myometrium, plays a vital role in implantation, immunoregulation, and angiogenesis, with the secretory endometrium producing PRL in response to MPA and E2, leading to ciliated cell formation and stromal cell decidualization (2, 20). HCG, secreted by trophoblasts in early pregnancy, influences decidual cells (21) and improves endometrial thickness and receptivity(22). The introduction of hCG has been shown to elevate the expression of critical factors associated with endometrial receptivity, such as endocytosis proteins, hypoxia-inducible factor 1 (HIF1), chemokines, and glycodelin (50). HPL contributes to the development and function of uterine glands(23). Consequently, the additional supplementation of PRL, hCG and HPL in our system augmented hormone responsiveness and receptivity, while promoting cell differentiation, ultimately yielding a model more representative of the receptive endometrium.

Earlier version: Here, our WOI organoids display specific features of in vivo receptive endometrium. They show similarly upregulated energy and lipid metabolism to the in vivo implantation window. Enhanced energy metabolism involves monocarboxylic acid metabolism and mitochondrial oxidative phosphorylation. Monocarboxylic acids, exemplified by lactate, pyruvate, and ketone bodies, are vital metabolites in most mammalian cells. During the implantation window, elevated lactate levels mobilize endometrial monocarboxylic acid metabolism and function as a pregnancy-related signal, stimulating ATP secretion in the epithelium. Subsequently, ATP, the primary product of energy metabolism, induces neighboring epithelial cells to release IL8, promoting decidualization of stromal cells (50). Lipid metabolism, responsible for energy storage, signal transduction, cell proliferation, apoptosis, and membrane trafficking, plays a crucial role in endometrial receptivity and implantation, although its mechanisms remain unclear (51–53). Thus, the WOI organoids build a metabolic framework for simulating in vivo implantation window.

Revised version: Here, our WOI organoids exhibit distinct characteristics resembling those of the in vivo receptive endometrium, including heightened energy and lipid metabolism. Enhanced energy metabolism involves monocarboxylic acid metabolism and mitochondrial oxidative phosphorylation. Monocarboxylic acids, exemplified by lactate, pyruvate, and ketone bodies, are vital metabolites in most mammalian cells. During the implantation window, elevated lactate levels mobilize endometrial monocarboxylic acid metabolism and function as a pregnancy-related signal, stimulating secretion in the epithelium. Subsequently, ATP, the primary product of energy metabolism, induces neighboring epithelial cells to release IL8, promoting decidualization of stromal cells (51). It was also observed that WOI organoids indeed produced much more ATP and IL8 than CTRL and SEC organoids (Fig.S3L). Lipid metabolism, responsible for energy storage, signal transduction, cell proliferation, apoptosis, and membrane trafficking, plays a crucial role in endometrial receptivity and implantation, although the precise mechanisms are not fully understood (52–54). Thus, the WOI organoids possessed metabolic characteristics of in vivo implantation window.

Earlier version: Functional alterations in secretory epithelium and EMT-derived stromal cells at the single-cell level also occur during the implantation window and warrant attention. The secretory epithelium responds to hypoxia through the HIF1 signaling pathway during the implantation window. HIF1A maintains oxygen homeostasis by increasing VEGF levels throughout the menstrual cycle, while HIF2A promotes angiogenesis for embryo implantation (20, 54). Angiogenesis and vascular remodeling support functional endometrial layer repair, endometrial thickening, receptivity, and embryo implantation. Importantly, EMT, a signature transition to the receptive endometrium, occurs in the proliferative epithelium and EMT-derived stromal cells and is involved in opening the epithelial layer of endometrium(38), during which epithelial cells lose polarity and cell-to-cell junctions, stromal cells acquire the migratory and invasive abilities (55). These implantation window associated cellular characteristics are reproduced in the WOI organoids, which confirms that they are currently the closest model to the in vivo state. The WOI organoids will allow us to explore the underlying mechanisms and advance the field of embryo implantation research.

Revised version: Functional alterations in secretory epithelium and EMT-derived stromal cells at the single-cell level also occur during the implantation window and warrant further investigation. The secretory epithelium responds to hypoxia through the HIF1 signaling pathway during the implantation window. HIF1A maintains oxygen homeostasis by increasing VEGF levels throughout the menstrual cycle, while HIF2A promotes angiogenesis for embryo implantation (25, 55). Angiogenesis and vascular remodeling support functional endometrial layer repair, endometrial thickening, receptivity, and embryo implantation. Importantly, EMT, a signature transition to the receptive endometrium, occurs in the proliferative epithelium and EMT-derived stromal cells and facilitates opening the epithelial layer of endometrium(43), during which epithelial cells lose polarity and cell-to-cell junctions, stromal cells acquire the migratory and invasive abilities (56). These implantation window associated cellular characteristics are partly reproduced in the WOI organoids, which confirms that they are closer to the in vivo state. The WOI organoids will allow us to explore the underlying mechanisms and advance the field of embryo implantation research.

Earlier version: However, our WOI endometrial organoids also have some limitations. We observed a decreased abundance of immune cells, primarily T cells characteristically expressing TSPYL2, RSRP1 and CD44, in WOI organoids, suggesting a possible immunoregulation in preparation for embryo implantation. But the immune cell composition of organoids varies from that of the in vivo endometrium. Although the organoids contain epithelial, stromal and immune cells, making them the closest model to the in vivo endometrium currently available, it is undeniable that they cannot perfectly replicate the in vivo state. The in vivo endometrium consists of functional and basal layers with more abundant cell subtypes than organoids due to the superior regulation by systemic hormones and nutrients. Moreover, the endometrial organoids lack the mechanical properties of the myometrium that impact in vivo endometrium (56). We are looking forward to constructing the whole uterus in vitro with stem cell induction, 3D printing technology, and microfluidic systems.

Revised version: However, our WOI endometrial organoids also exhibit some limitations. Firstly, a decreased abundance of immune cells, particularly T cells expressing TSPYL2, RSRP1, and CD44, was observed in WOI organoids, indicating potential immunoregulation of embryo implantation. However, the composition of immune cells in organoids differs from that of the in vivo endometrium. Specifically, the proportion of WBCs in organoids was found to be lower than that in the endometrium. Additionally, the proportions of T cells and macrophages in organoids were approximately 2% to 3% and 1% (Figure 2D), respectively, compared to 7% to 8% lymphocytes and 0.6% to 0.7% macrophages in the endometrium(1). Furthermore, during the implantation window, T cells in WOI organoids decreased (Figure S2F), while T cells in the endometrium increased(1). Secondly, it is undeniable that the organoids cannot perfectly replicate the in vivo endometrium, which comprises functional and basal layers with a greater abundance of cell subtypes, under superior regulation by systemic hormones and nutrients. Thirdly, the endometrial organoids lack the mechanical properties of the myometrium, which have an impact on the in vivo endometrium (57). We are looking forward to constructing the whole uterus in vitro with stem cell induction, 3D printing technology, and microfluidic systems.

Earlier version: In summary, we developed a receptive endometrial organoid that mimics the features of in vivo implantation-window endometrium features (Fig. 5). Our human endometrial organoid comprises epithelium, stromal cells, and immune cells, ensuring hormone responsiveness and endometrial secretion function while recapitulating aspects such as decidualization, ECM remodeling, pinopode formation, cilia generation, EMT and metabolism. This receptive endometrial organoid serves as a platform to investigate endometrial physiology, maternal-fetal interactions, and functional or organic diseases, with potential for functional applications and clinical translation.

Revised version: In summary, we developed a receptive endometrial organoid that mimics the features of in vivo implantation-window endometrium (Fig. 5). Our human endometrial organoid comprises epithelial cells, stromal cells, and immune cells, ensuring hormone responsiveness and endometrial secretion function while recapitulating aspects such as decidualization, ECM remodeling, pinopode formation, cilia generation, EMT and metabolism. WOI organoids exhibit differentiation from SEC organoids towards the implantation window. The SEC organoids are suitable for studying endometrial secretory phase and hormone reactivity. This receptive endometrial organoid serves as a platform to investigate peri-implantation endometrial physiology and pathology, maternal-fetal interactions, with potential practical applications and clinical translation.

Earlier version: Schematic diagram displaying the establishment and validation of receptive endometrial organoids, and summarizing the characteristic biological events of endometrial implantation window.

Revised version: Schematic diagram displaying the establishment and validation of receptive endometrial organoids, and summarizing the characteristic biological events of implantation-window endometrium.

Materials and Methods

Establishment of endometrial organoids

All experiments involving human subjects followed medical ethical principles and the Declaration of Helsinki, and were approved by the Ethics Committee of Qilu Hospital of Shandong University. Informed consents were obtained from patients. Females of reproductive age who received hysterectomy for benign diseases were selected for this study. Clinical information of patients providing endometrial tissue was listed in Dataset S1. Experimental procedures are presented in detail in the Supporting Information.

Hormone treatment of endometrial organoids

Earlier version: The hormone regimen for inducing endometrial secretory phase, as described by Thomas E. Spencer et al (7), consisted of estradiol for 2 days followed by a combination of estradiol (Sigma E2758), medroxyprogesterone acetate (Selleck S2567) and N6,2′-O-dibutyryladenosine 3′,5′-cyclic monophosphate sodium salt (cAMP) (Sigma D0627) for 6 days.

Revised version: The hormone regimen for inducing endometrial secretory phase, as described by Thomas E. Spencer et al (8), consisted of estradiol for 2 days followed by a combination of estradiol (Sigma E2758), medroxyprogesterone acetate (Selleck S2567) and N6,2′-O-dibutyryladenosine 3′,5′-cyclic monophosphate sodium salt (cAMP) (Sigma D0627) for 6 days.

Earlier version: To simulate the endometrium of implantation window, we developed a model by incorporating various pregnancy-related hormones. The hormone regimen included estradiol for the first two days, followed by a combination of estradiol, medroxyprogesterone acetate, cAMP, Human Chorionic Gonadotropin (HCG) (Livzon Pharmaceutical Group Inc), Human Placental Lactogen (HPL) (R&D Systems 5757-PL), and prolactin (Peprotech 100-07) for 6 days. The CTRL group was cultured in ExM without the addition of hormones. (Fig. 1A)

Revised version: To simulate the endometrium of implantation window, we developed a model by incorporating various pregnancy-related hormones. The hormone regimen included estradiol for the first two days, followed by a combination of estradiol, medroxyprogesterone acetate, cAMP, Human Chorionic Gonadotropin (HCG) (Livzon Pharmaceutical Group Inc), Human Placental Lactogen (HPL) (R&D Systems 5757-PL), and prolactin (Peprotech 100-07) for 6 days. The CTRL group was cultured in ExM without the addition of hormones. (Fig. 1A) (Table S2)

Supporting information